# COMP SCI 2009 Programming for IT - W05 Practical

## Set Covering
As mentioned in the lecture, the **set cover problem** is an NP-complete problem. In this task, you will solve this question using both exhausive search (brute force) and greedy apporach. 

The problem is defined as follow:

Given a `list` of `sets` where each `set` contains integers, please find the minimum number of sets that cover the integers from `1` to `10`

For example, given the following three sets:
1. `[1, 2, 3, 4, 5]`
1. `[2, 4, 6, 8]`
1. `[6, 7, 8, 9, 10]`

Your code should find that `set1` and `set3` will cover all the integers from `1` to `10`.


## Exercise 1: Set cover - brute force 

First, implement the brute force approach that tests `all` combinations. You can assume there will be a solution. To simplify your code a bit, you can use the `itertools`. Please read its document and examples below:

- https://www.geeksforgeeks.org/itertools-combinations-module-python-print-possible-combinations/
- https://docs.python.org/3/library/itertools.html

<!-- grade_id: ans-1, points: 0, grade: false, solution: true -->
```python
import itertools

def brute_force_set(li):
    # INSERT YOUR CODE BELOW
    # use itertools.combinations
    # ~ 8 lines
    ### BEGIN SOLUTION
    for r in range(0, len(li)+1):
        for subset in itertools.combinations(li, r):
            s = {}
            for sset in subset:
                for e in sset:
                    s[e] = ''
            if(len(s) == 10):
                return subset
    ### END SOLUTION

# =======================
# EXAMPLE TESTING CASES
set_list1 = [ # first test input
    [1, 2, 3],
    [4, 5, 6],
    [7, 8, 9, 10],
    [1, 3, 7],
    [2, 4, 6, 8]
]

set_list2 = [ # second test input
    [4, 2, 5],
    [1, 7],
    [2, 4, 8, 9, 10],
    [5, 10],
    [3, 5, 6],
    [2, 6],
    [1, 8, 9, 10]
]
            
ss = brute_force_set(set_list1)
print(ss) # ([1, 2, 3], [4, 5, 6], [7, 8, 9, 10])
print('---')
ss = brute_force_set(set_list2)
print(ss) # ([1, 7], [2, 4, 8, 9, 10], [3, 5, 6])
```

Output:
```
([1, 2, 3], [4, 5, 6], [7, 8, 9, 10])
---
([1, 7], [2, 4, 8, 9, 10], [3, 5, 6])
```

<!-- grade_id: hiddentest1-1, points: 0.5, grade: true, solution: false -->
```python
# Testing Cell (Do NOT modify this cell)
### BEGIN HIDDEN TESTS
s = [[1, 2, 3], [4, 5, 6, 7], [8, 9, 10], [3, 5]]
ss = brute_force_set(s)
assert ss[-1] == [8, 9, 10], "brute_force_set failed a test"
### END HIDDEN TESTS
```

<!-- grade_id: hiddentest1-2, points: 0.5, grade: true, solution: false -->
```python
# Testing Cell (Do NOT modify this cell)
### BEGIN HIDDEN TESTS
s = [[1, 2, 3, 4, 5, 6, 7, 8], [2, 4, 6], [9], [10], [8]]
ss = brute_force_set(s)
assert ss[-1] == [10], "brute_force_set failed a test"
### END HIDDEN TESTS
```

## Exercise 2: Set cover - greedy



In this exercise, please implement a greedy algorithm that **for each step, takes a set that contains the most uncovered integers** until it covers all integers from `1` to `10`.

<!-- grade_id: ans-2, points: 0, grade: false, solution: true -->
```python
import copy
def greedy_set(list_sets):
    sets = list_sets.copy()
    ans = [] # your answer, should be a list of sets
    covered = {}

    kk = 0
    while True:
        # INSERT YOUR CODE BELOW
        # Step 1: search for the set with most uncovered integers
        # Step 2: add the set into ans
        # Step 3: check if ans already cover all integers
        # ~ 15 lines
        
        ### BEGIN SOLUTION
        # search for the set with most uncovered integers 
        max_c = 0
        max_s = {}

        for ss in sets:                
            # calculate number of uncovered integers
            c = 0
            for i in ss:
                if i not in covered:
                    c += 1         
                    
            # update current best solution
            if c > max_c:
                max_s = ss
                max_c = c
#                 print(f'{max_s} {max_c}')
            
        # add into the set
#         print(f'append {max_s}')
        ans.append(max_s)
        sets.remove(max_s)
        
        for i in max_s:
            covered[i] = 0

        # check if all numbers are included
        if(len(covered) == 10):
            return ans
        ### END SOLUTION

        
        # BELOW 3 LINES ARE TO AVOID INFINITE LOOP FOR YOU
        kk += 1
        if kk > 10:
            break

# =======================
# EXAMPLE TESTING CASES
set_list1 = [ # first test input
    [1, 2, 3],
    [4, 5, 6],
    [7, 8, 9, 10],
    [1, 3, 7],
    [2, 4, 6, 8]
]

set_list2 = [ # second test input
    [4, 2, 5],
    [1, 7],
    [2, 4, 8, 9, 10],
    [5, 10],
    [3, 5, 6],
    [2, 6],
    [1, 8, 9, 10]
]

s = greedy_set(set_list1)
print(s) # [[7, 8, 9, 10], [1, 2, 3], [4, 5, 6]]
print('---')
s = greedy_set(set_list2)
print(s) # [[2, 4, 8, 9, 10], [3, 5, 6], [1, 7]]
```

Output:
```
[[7, 8, 9, 10], [1, 2, 3], [4, 5, 6]]
---
[[2, 4, 8, 9, 10], [3, 5, 6], [1, 7]]
```

<!-- grade_id: hiddentest2-1, points: 0.5, grade: true, solution: false -->
```python
# Testing Cell (Do NOT modify this cell)
### BEGIN HIDDEN TESTS
s = [[1, 2, 3], [4, 5, 6, 7], [8, 9, 10], [3, 5]]
ss = greedy_set(s)
assert ss[-1] == [8, 9, 10], "greedy_set failed a test"
### END HIDDEN TESTS
```

<!-- grade_id: hiddentest2-2, points: 0.5, grade: true, solution: false -->
```python
# Testing Cell (Do NOT modify this cell)
### BEGIN HIDDEN TESTS
s = [[1, 2, 3, 4, 5, 6, 7, 8], [2, 4, 6], [9], [10], [8]]
ss = greedy_set(s)
assert ss[-1] == [10], "greedy_set failed a test"
### END HIDDEN TESTS
```

# Exercise 3 - Knapsack Problem

It’s Boxing Day and you are in a shopping contest! You have one shopping bag and you would like to fill your bags with items that have the largest sum of values. Please solve this problem using dynamic programming.

<!-- grade_id: ans-3, points: 0, grade: false, solution: true -->
```python
def knapsack_dp(bag_size, items):
    num_row = len(items)
    num_col = bag_size    
    grid = [[0]* (num_col+1) for _ in range(num_row)] # initialise the grid    
    prev_max = 0
    
    # INSERT YOUR CODE
    # Update the DP table (the variable grid)
    # ~ 15 lines
    ### BEGIN SOLUTION
    for i in range(0, num_row):
        w = items[i][1]
        v = items[i][2]
        
        for j in range(1, num_col+1):
            if(i != 0):
                prev_max = grid[i-1][j]            
            else:
                prev_max = 0
                
            if j-w >= 0:
                current = v + grid[i-1][j-w]                      
            else:
                current = 0
            grid[i][j] = max(prev_max, current)            
    ### END SOLUTION
    return grid

def display(grid):
    for r in grid:
        s = ' '
        for c in r:
            s += f'{c:4} '
        print(s)

# =======================
# EXAMPLE TESTING CASES
bag1 = 4 # the maximum weight for the bag
item_list1 = [
    ['guitar', 4, 3000],
    ['speaker', 1, 1500],
    ['laptop', 3, 2000]
]

bag2 = 7 # the maximum weight for the bag
item_list2 = [
    ['pasta', 1, 700],
    ['soups', 3, 300],
    ['pork', 3, 700],
    ['steak', 4, 1000],
    ['lamb', 3, 1200],
    ['Cheese', 1, 900],
    ['Mushroom', 1, 800],
]

g = knapsack_dp(bag1, item_list1)
display(g)
print('---')
g = knapsack_dp(bag2, item_list2)
display(g)

# output should be as below

#     0    0    0    0 3000 
#     0 1500 1500 1500 3000 
#     0 1500 1500 2000 3500 
# ---
#     0  700  700  700  700  700  700  700 
#     0  700  700  700 1000 1000 1000 1000 
#     0  700  700  700 1400 1400 1400 1700 
#     0  700  700  700 1400 1700 1700 1700 
#     0  700  700 1200 1900 1900 1900 2600 
#     0  900 1600 1600 2100 2800 2800 2800 
#     0  900 1700 2400 2400 2900 3600 3600 
```

Output:
```
0    0    0    0 3000 
    0 1500 1500 1500 3000 
    0 1500 1500 2000 3500 
---
    0  700  700  700  700  700  700  700 
    0  700  700  700 1000 1000 1000 1000 
    0  700  700  700 1400 1400 1400 1700 
    0  700  700  700 1400 1700 1700 1700 
    0  700  700 1200 1900 1900 1900 2600 
    0  900 1600 1600 2100 2800 2800 2800 
    0  900 1700 2400 2400 2900 3600 3600
```

<!-- grade_id: hiddentest3-1, points: 1, grade: true, solution: false -->
```python
# Testing Cell (Do NOT modify this cell)
### BEGIN HIDDEN TESTS
item_list1 = [['a', 4, 3000],['b', 1, 1500],['c', 3, 2000]]
g = knapsack_dp(5, item_list1)
assert g[-1][-1] == 4500, "knapsack_dp failed a test"
### END HIDDEN TESTS
```

<!-- grade_id: hiddentest3-2, points: 1, grade: true, solution: false -->
```python
# Testing Cell (Do NOT modify this cell)
### BEGIN HIDDEN TESTS
item_list1 = [['a', 2, 3000],['b', 1, 1500],['c', 1, 2000]]
g = knapsack_dp(5, item_list1)
assert g[-1][-1] == 6500, "knapsack_dp failed a test"
### END HIDDEN TESTS
```